---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: c5f30ed116
```



`@instructions`
ok here I'm





`@sct`
```{python}
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa#
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa#
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa#
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa#
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: b2b6771e1a
```





`@pre_exercise_code`
```{python}
import geopandas
```








---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise

xp: 50

key: 79d151c4cb
```

this is assignment text


`@hint`
hint





`@possible_answers`
- 1
- 2
- [3]

`@feedbacks`
- no
- no
- yes!




---
## Insert exercise title here

```yaml
type: TabExercise

xp: 100

key: 96ca222bd3
```













***



```yaml
type: NormalExercise

xp: 50

key: 74dc6062f8
```



`@instructions`
d;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklas

`@hint`
d;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklasd;asdklas

`@pre_exercise_code`
```{undefined}
import math from pi
```
`@sample_code`
```{undefined}
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa#
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa#
```
`@solution`
```{undefined}
# here is the solition
```







***



```yaml
type: NormalExercise

xp: 50

key: 87f98074aa
```



`@instructions`


`@hint`





`@sct`
```{undefined}
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa#
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa
# sa# sa# sa# # sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa#
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# sa
# sa# sa# sa# 
# sa# sa# sa# sa
# sa# sa# sa#
```






---
## Insert exercise title here

```yaml
type: VideoExercise

xp: 50

key: cba5fc79cf
```

`@projector_key`
4b4fe1e445f9b3a98c03509bc8e0e781

---
## Insert exercise title here

```yaml
type: TabExercise

xp: 100

key: 7a8a24817e
```












