---
  title: "Test"
  description: "Test"
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 53f45190f6   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```

